---
title       : Insert the chapter title here
description : Insert the chapter description here

---
## Test

```yaml
type: NormalExercise
key: d2ac5807ea
lang: python
xp: 100
skills: 2
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```



---
## Test 2

```yaml
type: MultipleChoiceExercise
key: 6a4dd769f4
lang: python
xp: 50
skills: 2
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```

---
## Normal one

```yaml
type: NormalExercise
key: 13550aaa31
lang: python
xp: 100
skills: 2
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}
Ex().has_code('nothing!')
```
